# High Pass Filter Using the Spectral Reversal Technique
We can create a high pass filter by using as reference a low pass filter and a technique called **Spectral Reversal**. For this notebook we will use the *Windowed-Sinc Filters* Notebook results, which are pickled in an serialized object called `save_data.pickle`.

<font color="blue">Note: In order to run this Jupyter Notebook you must create a file named `aux_functions.py` inside the folder containing this Notebook with the following functions `get_fourier`, `shifted_sinc_function`,  `hamming_window` and `blackman_window` which were developed previously.</font> 

```python
import numpy as np
import matplotlib.pyplot as plt

import pickle

from aux_functions import get_fourier
from aux_plots import plot_frequency_response
```

## Load the data and the Low-Pass Filter Coefficients
First we load the ECG data that we used on the *Windowed-Sinc Filters* Notebook, and then we loaded the low pass filter stored in a piclke file.

```python
with open('save_data.pickle', 'rb') as f:
    data = pickle.load(f)
    
ecg = np.array(data['ecg'])
low_pass = np.array(data['low_pass'])
low_pass = low_pass/np.sum(low_pass)
fft_low_pass = np.array(data['fft_low_pass'])
```

## Create a High-Pass Filter
To generate the high pass filter, we use the **Sprectral Reversal** methond, which consist of multiplying the low pass filter response $h_{lp}[n]$ with $(-1)^{n}$. 

Therefore, the high pass filter response is given by:
$$h_{hp}[n] = h_{lp}[n](-1)^{n}$$

Implement a high-pass filter function called `spectral_reversal` wich takes as input an array of low-pass filter coefficients `h_lp` and returns an array of high-pass filter coefficient using the equation described before.

<!-- grade_id: cell-0f23ec469e28784b, points: 0, grade: false, solution: true -->
```python
def spectral_reversal(h_lp):
    """
    Function that performs the spectral reversal of a low pass filter with
    coefficients h_lp and returns the high pass coefficients h_hp.
  
    Parameters: 
    h_lp (array): Array of numbers representing the low pass filter coefficients.
  
    Returns: 
    numpy array: Array of numbers representing the high pass filter coefficients.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-58145ea77af49c41, points: 0, grade: false, solution: true -->
```python
# Use the spectral_reversal function on low pass filter to obtain a high pass filter.
# The coefficients of the low pass filter are stored in the low_pass variable.
# Store the result in a variable named high_pass.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-af4a3e31fe4775c1, points: 50, grade: true, solution: false -->
```python
with open('high_pass.pkl', 'rb') as file:
    high_pass_pkl = pickle.load(file)
    

assert np.allclose(high_pass_pkl, high_pass, atol=0.01)
```

<!-- grade_id: cell-ce0b045e04257243, points: 0, grade: false, solution: true -->
```python
# Find the Fourier Transform for the low pass filter and assign:
# The magnitude to dft_low_pass_magnitude
# The normalized frequency to dft_low_pass_freq
# YOUR CODE HERE
raise NotImplementedError()

# Find the Fourier Transform for the high pass filter and assign:
# The magnitude to dft_high_pass_magnitude
# The normalized frequency to dft_high_pass_freq
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-988136e4e107b7a1, points: 25, grade: true, solution: false -->
```python
with open('high_pass_dft.pkl', 'rb') as file:
    dft_low_pass_magnitude_pkl, dft_low_pass_freq_pkl, \
    dft_high_pass_magnitude_pkl, dft_high_pass_freq_pkl = pickle.load(file)
    
    
assert np.allclose(dft_low_pass_magnitude_pkl, dft_low_pass_magnitude, atol=0.01)
assert np.allclose(dft_low_pass_freq_pkl, dft_low_pass_freq, atol=0.01)
assert np.allclose(dft_high_pass_magnitude_pkl, dft_high_pass_magnitude, atol=0.01)
assert np.allclose(dft_high_pass_freq_pkl, dft_high_pass_freq, atol=0.01)


plt.rcParams["figure.figsize"] = (15,10)

plt.subplot(2,2,1)
#plt.stem(low_pass, markerfmt='.', use_line_collection=True)
plt.plot(low_pass)
plt.title('Low Pass Filter')
plt.grid('on')

plt.subplot(2,2,2)
#plt.stem(high_pass, markerfmt='.', use_line_collection=True)
plt.plot(high_pass)
plt.title('High Pass Filter')
plt.grid('on')

plt.subplot(2,2,3)
dft_low_pass_magnitude_reshape = dft_low_pass_magnitude
plot_frequency_response(dft_low_pass_magnitude_reshape.reshape(-1,1), 
                               dft_low_pass_freq, 
                               title='Low Pass Filter Response')

plt.subplot(2,2,4)
dft_high_pass_magnitude_reshape = dft_high_pass_magnitude
plot_frequency_response(dft_high_pass_magnitude_reshape.reshape(-1,1), 
                               dft_high_pass_freq, 
                               title='High Pass Filter Response');
```

As you can see, the frequency response of the high pass filter is a shifted version of the low-pass filter. Specifically, it is shifted $\pi$ rads or 0.5.

<!-- grade_id: cell-811fa3713096a4b3, points: 0, grade: false, solution: true -->
```python
# Filter the ECG signal with the low pass filter. For this, use the np.convolve function.
# Set the mode parameter to 'same'
# YOUR CODE HERE
raise NotImplementedError()

# Filter the ECG signal with the high pass filter. For this, use the np.convolve function.
# Set the mode parameter to 'same'
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fa0b2ae33f3b6b1d, points: 25, grade: true, solution: false -->
```python
with open('convolve.pkl', 'rb') as file:
    low_pass_ecg_pkl, high_pass_ecg_pkl = pickle.load(file)
    
assert np.allclose(low_pass_ecg_pkl, low_pass_ecg, atol=0.01)
assert np.allclose(high_pass_ecg_pkl, high_pass_ecg, atol=0.01)

plt.rcParams["figure.figsize"] = (15,10)

plt.subplot(2,2,1)
plt.plot(low_pass_ecg)
plt.title('Low Pass ECG')
plt.grid('on')
plt.xlabel('Samples')
plt.ylabel('Amplitude')

plt.subplot(2,2,2)
plt.plot(high_pass_ecg)
plt.title('High Pass ECG')
plt.grid('on')
plt.xlabel('Samples')
plt.ylabel('Amplitude')

plt.subplot(2,1,2)
plt.plot(ecg)
plt.title('ECG Signal')
plt.grid('on')
plt.xlabel('Samples')
plt.ylabel('Amplitude')
plt.show()
```

## Why Does Spectral Reversal Work?

The spectral reversal technique is based on the so-called shift theorem of the Fourier transform. Formulated for the discrete case, the shift theorem says that, for a Fourier transform pair $x[n]⟷X[k]$, a shift by $s$ samples in the frequency domain is equivalent with multiplying by a complex exponential in the time domain, as

$$x[n]e^{j2πns/N}⟷X[k−s]$$

In general, the result of doing this will be complex. This is not a problem, but let’s investigate in which circumstances the result is not complex. This is may be easier to see with the complex exponential rewritten using Euler’s Identity,

$$e^{j2πns/N}=\cos{(2πns/N)}+j\sin{(2πns/N)}$$

This expression is real if the sine term is zero, which is the case if the angle is zero or $\pi$

We get this result if we shift by $s=N/2$ samples, because then the angles become $2πns/N=nπ$. Of course, in that case the complete expression becomes $\cos{(nπ)}$, which is exactly the sequence $1,−1,1,−1,…$.

#### Reference:
* https://tomroelandts.com/articles/spectral-reversal-to-create-a-high-pass-filter