# Digital Filters
In this part we will introduce the two types of digital filters and how the z-transform is used for calculating their frequency response, and their zeros and poles.


## The Z-Transform

Just as analog filters are designed using the Laplace transform, recursive digital filters are developed with a parallel technique called the z-transform. The overall strategy of these two transforms is the same: probe the impulse response with sinusoids and exponentials to find the system's poles and zeros. The Laplace transform deals with differential equations, the s-domain, and the s-plane. Correspondingly, the z-transform deals with difference equations, the z-domain, and the z-plane. However, the two techniques are not a mirror image of each other; the s-plane is arranged in a rectangular coordinate system, while the z-plane uses a polar format. 

The standard form of the z-transform is given by:
$$X(z) = \sum \limits^{\infty}_{n=-\infty} x[n]z^{-n} $$

where $z=re^{j\omega}$ and $r=e^{\sigma}$.

The utility to have the z-transform resides in the implementation of recursive filters, which uses *recursion coefficients*. To analyze these systems in the z-domain, we must be able to convert these recursion coefficients into the z-domain *transfer function*, and back again. Defining the z-transform in this manner ($r^n$ and $z$) provides the simplest means of moving between these two important representations. In fact, defining the z-domain in this way makes it trivial to move from one representation to the other.

## Stability and Relation with the Fourier Transform

In a continous system, if a pole occupies the right half of the s-plane (by the use of the Laplace transform) the system will be unstable. In the case of a discrete system, the system will be unstable if the poles are outside the unit circle in the z-plane. When the time domain signal is completly real, the upper and lower halves of the z-plane are mirror images of each other, just as with the s-domain.

A continous sinusoid can have any frequency between DC and infinity, a discrete sinusoid can only have a frequency between DC and one-half of the sampling rate. That is, frequency must be between 0 and 0.5, or between 0 and $\pi$ when expressed as a natural frequency.

The Laplace transform, evaluated at $\sigma = 0$, is equal to the Fourier transform. Evaluating the z-transform at $r=1$ is equal to the Discrete Time Fourier Transform (DTFT).

## Analysis of Recursive Systems

A recursive filter is described by a **difference equation**:

$$ y[n] = a_0 x[n] + a_1 x[n-1] + a_2 x[n-2] + a_3 x[n-3] + \cdot \cdot \cdot + b_1 y[n-1] + b_2 y[n-2] + b_3 y[n-3] + \cdot \cdot \cdot $$ <br>
where $x$ and $y$ are the input and output signals respectively, and $a$ and $b$ terms are the **recursion coefficients**.

Applying the z-transform to difference equation gives us:

$$ Y[z] = a_0 X(z) + a_1 X(z) z^{-1} + a_2 X(z) z^{-2} + a_3 X(z) z^{-3} + \cdot \cdot \cdot + b_1 Y(z)z^{-1} + b_2 Y(z)z^{-2} + b_3 Y(z)z^{-3} + \cdot \cdot \cdot $$


and the **system's transfer function** is:

$$H[z] = \frac{a_0 + a_1z^{-1} + a_2z^{-2} + a_3z^{-3} + \cdot \cdot \cdot}{1 - b_1z^{-1} - b_2z^{-2} - b_3z^{-3}} $$


## IIR Filters and FIR Filters

The impulse response or the frequency response classify digital filters. If the impulse response of the filter falls to zero after a finite period of time, it is an FIR (Finite Impulse Response) filter. However, if the impulse response exists indefinitely, it is an IIR (Infinite Impulse Response) filter. How the output values are calculated determines whether the impulse response of a digital filter falls to zero after a finite period of time. For FIR filters the output values depend on the current and the previous input values, whereas for IIR filters the output values also depend on the previous output values.

For a **FIR filter** we have the following difference equation

$$ y[n] = a_0 x[n] + a_1 x[n-1] + a_2 x[n-2] + \cdot \cdot \cdot + a_k x[n-k] $$ 

and the system's transfer function

$$H[z] = a_0 + a_1z^{-1} + a_2z^{-2} + \cdot \cdot \cdot + a_kz^{-k} $$

<br>

For a **IIR filter** we have the following difference equation:

$$ y[n] = a_0 x[n] + a_1 x[n-1] + a_2 x[n-2] + a_3 x[n-3] + \cdot \cdot \cdot + b_1 y[n-1] + b_2 y[n-2] + b_3 y[n-3] + \cdot \cdot \cdot $$ 

and the system's transfer function:

$$H[z] = \frac{a_0 + a_1z^{-1} + a_2z^{-2} + a_3z^{-3} + \cdot \cdot \cdot}{1 - b_1z^{-1} - b_2z^{-2} - b_3z^{-3}} $$

<br>

### Advantages and Disadvantages of FIR and IIR Filters

The advantage of IIR filters over FIR filters is that **IIR filters usually require fewer coefficients** to execute similar filtering operations, that IIR filters work faster, and require less memory space.

The disadvantage of **IIR filters** is the **nonlinear phase response**. IIR filters are well suited for applications that require no phase information, for example, for monitoring the signal amplitudes. FIR filters are better suited for applications that require a linear phase response.

## Implementation:
In the following section you will see the frequency response and, the poles and zeros of a filter given the coefficients of system's transfer function.

We will work with the following transfer function:

$$H[z] = \frac{1.000 -1.414z + 1.000z^{2} }{0.810 -1.273z +1.000z^{2}} $$

Note that our transfer function has the form:

$$H[z] = \frac{c_0 + c_1z + c_2z^{2} + c_3z^{3} + \cdot \cdot \cdot}{d_0 + d_1z + d_2z^{2} + d_3z^{3}+ \cdot \cdot \cdot} $$

```python
import matplotlib.pyplot as plt
import numpy as np
import scipy.signal
import pickle

from aux_plots import plot_zeros_poles
```

### Create a `filter_frequency_response` function
First you need to create a function called `filter_frequency_response` which takes as arguments the numerator `c` coefficients and the denominator `d` coefficients of the $H[z]$ function described before. In this case `w` is the range of the frequency sweep and by default has a step size of 0.1. What your function needs to do is to return the values of $H[z]$ for all the specified `w` values. (Remember that $z = e^{j\omega}$)

<!-- grade_id: cell-67eeee268b577bed, points: 0, grade: false, solution: true -->
```python
def filter_frequency_response(c, d, w = np.arange(0, np.pi, 0.1)):
    """ 
    Function that generates the frequency response of a digital filter given the coeficients of 
    polynomials c0 + c_1*x + c_2*x^2 + ... and d0 + d_1*x + d_2*x^2 + ...
    This function evaluates the transfer function 
    H(x)=(c0 + c_1*x + c_2*x^2 + ...)/(d0 + d_1*x + d_2*x^2 + ...) where x is an element in vector w.
  
    Parameters: 
    w (numpy array): Array of natural frequency values.
    c (numpy array): Array of recursion coefficients c.
    d (numpy array): Array of recursion coefficients d.
  
    Returns: 
    numpy array: Returns filter response.
  
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3f37f7dd9cd00949, points: 60, grade: true, solution: false -->
```python
with open('filter_frequency_response.pkl', 'rb') as file:
    H_w_pkl = pickle.load(file)
    

c = np.array([1, -1.414, 1])
d = np.array([0.810, -1.273, 1])
w = np.arange(0, np.pi, 0.01)

H_w = filter_frequency_response(c,d,w)


plt.rcParams["figure.figsize"] = (15,5)

plt.plot(w/(2*max(w)), 20*np.log(np.absolute(H_w)), label='Calculated')
plt.plot(w/(2*max(w)), 20*np.log(np.absolute(H_w_pkl)), label='Expected')
plt.title('Frequency Response')
plt.xlabel('frequency')
plt.ylabel('dB')
plt.grid('on')
plt.legend()
plt.show()

assert np.isclose(H_w, H_w_pkl, atol=0.01).all()
```

### Create a `zeros_poles_gain` function
Now you will need to create a function that returns the zeros, poles and gain of a given transfer function. Again, it's inputs are the numerator `c` coefficients and the denominator `d` coefficients of $H[z]$. For implementing this function you can use `scipy.signal.tf2zpk` according to your needs.

<!-- grade_id: cell-67eeee268b5774, points: 0, grade: false, solution: true -->
```python
def zeros_poles_gain(c, d):
    """ 
    Function that calculates the zeros, poles and gain of a given transfer function which consists of
    the coeficients of polynomials c0 + c_1*x + c_2*x^2 + ... and d0 + d_1*x + d_2*x^2 + ...
  
    Parameters: 
    c (numpy array): Array of recursion coefficients c.
    d (numpy array): Array of recursion coefficients d.
  
    Returns: z,p,g
    z (numpy array): Zeros of transfer function.
    p (numpy array): Poles of transfer function.
    g (numpy array): Gain of transfer function.
  
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-597983b03a94100e, points: 20, grade: true, solution: false -->
```python
z, p, g = zeros_poles_gain(c, d)

plt.rcParams["figure.figsize"] = (15,5)
plot_zeros_poles(z, p)
plt.title('Zeros and Poles')
plt.xlabel('sample')
plt.show()

with open('zeros_poles_gain.pkl', 'rb') as file:
    z_pkl, p_pkl, g_pkl  = pickle.load(file)
    
assert np.isclose(z, z_pkl, atol=0.01).all()
assert np.isclose(p, p_pkl, atol=0.01).all()
assert np.isclose(g, g_pkl, atol=0.01).all()
```

### See all results
Finally you will see all the results together, including the Inverse Fourier Transform of $H[\omega]$, to better understand your filter.

<!-- grade_id: cell-821df22e47c5a702, points: 0, grade: false, solution: true -->
```python
c = np.array([1, -1.414, 1])
d = np.array([0.810, -1.273, 1])
w = np.arange(0, np.pi, 0.01)

H_w = filter_frequency_response(c,d,w)
# Find the Inverse Fourier Transform of H_w and assign it to idft
# YOUR CODE HERE
raise NotImplementedError()
z, p, g = zeros_poles_gain(c, d)
```

<!-- grade_id: cell-60112b5f80f2a81d, points: 20, grade: true, solution: false -->
```python
with open('idft.pkl', 'rb') as file:
    idft_pkl = pickle.load(file)
    
assert np.isclose(idft, idft_pkl, atol=0.01).all()


plt.rcParams["figure.figsize"] = (15,5)

plt.subplot(1, 3, 1)
plt.plot(w/(2*max(w)), 20*np.log(np.absolute(H_w)))
plt.title('Frequency Response')
plt.xlabel('frequency')
plt.ylabel('dB')
plt.grid('on')

plt.subplot(1, 3, 2)
plt.stem(np.real(idft[0:25]), use_line_collection=True)
plt.title('Impulse Response')
plt.xlabel('sample')
plt.grid('on')

plot_zeros_poles(z, p)
plt.title('Zeros and Poles')
plt.xlabel('sample')
plt.show()
```

#### Reference
* https://www.dspguide.com/ch33.htm
* https://docs.scipy.org/doc/scipy/reference/generated/scipy.signal.tf2zpk.html
* https://zone.ni.com/reference/en-XX/help/370858P-01/genmaths/genmaths/calc_filterfir_iir/
* https://www.minidsp.com/applications/dsp-basics/fir-vs-iir-filtering
*  https://dsp.stackexchange.com/questions/41642/filter-design-with-zero-pole-placement-method

```python

```