# Introduction
This Jupyter Notebook will be a warm up for the course work, and it will teach you how to use the material during this semester.

Usually there are two types of cells that you will encounter:
1. Code based cells:
    * Read only cell
    * Complete assigment cell
    * Auto test and grading cell


2. Markdown text cells
    * Read only cell
    * Complete assigment cell

## 1.1 Code based cells: Read only cell
An example of a read only code base cell is the following one, where different libraries are imported. Usually you will see that *NumPy*, *matplotlib* and *pickle* are among the libraries that we use throught the course.

```python
import numpy as np
import matplotlib.pyplot as plt
import pickle
```

## 1.2 Code based cells: Complete assigment cell
This type of cells are for you to complete an assigment that later will be graded. Usually some extra information is provided for you to complete a task. The following is an example of this type of cell.

<!-- grade_id: cell-91aa5ce7f0ca9763, points: 0, grade: false, solution: true -->
```python
# Create a function named `distance` that calculates the euclidian distance between two vectors.

def distance(vect1, vect2):
    """
    Function that calculates the distance between vector vect1 and vect2.
    Parameters: 
    vect1 (numpy array): Array of numbers representing a vector.
    vect2 (numpy array): Array of numbers representing a vector.
  
    Returns: 
    distance (float): Returns distance betwee two vectors
    
    """
    # YOUR CODE HERE
    distance = np.sqrt(np.sum(np.square(vect1 - vect2)))
    return distance
```

## 1.3 Code based cells: Auto test and grading cell
After you completed you solution, you can test and grade your code. For this, code based test and grading cells are used. The following cell shows how this cells work.

<!-- grade_id: cell-40438960591853ff, points: 100, grade: true, solution: false -->
```python
vect1 = np.array([5, 5, 5])
vect2 = np.array([4, 4, 4])

# Test using calculation
assert (distance(vect1, vect2) == np.sqrt(3))

# Test using pickle
with open('test.pkl', 'rb') as file:
    dis_pkl = pickle.load(file)
    
assert np.allclose(distance(vect1, vect2), dis_pkl, atol=0.0001)
```

## 2.1 Markdown text cells: Read only cell
This is an example of a read only markdown cell. Almost all of the text you will see in these Notebooks uses read only cells.

## 2.2 Markdown text cells: Complete assigment cell
This type of cell expects that you input some text information using markdown format. 

Answer in the following cell, what is your name?

<!-- grade_id: cell-af404929162a46a1, points: 0, grade: true, solution: true -->
Manuel Rios

## 3. NumPy
NumPy is the fundamental package for scientific computing in Python. It is a Python library that provides a multidimensional array object, various derived objects, and an assortment of routines for fast operations on arrays.

## 3.1 Arrays

<!-- grade_id: cell-4ab15d61d169d99f, points: 0, grade: false, solution: true -->
```python
# Define two arrays
# YOUR CODE HERE
raise NotImplementedError()

# Operations
# Addition
# YOUR CODE HERE
raise NotImplementedError()

# Subtraction
# YOUR CODE HERE
raise NotImplementedError()

# Multiplication (element wise)
# YOUR CODE HERE
raise NotImplementedError()

# Division (element wise)
# YOUR CODE HERE
raise NotImplementedError()
```

## 3.2 Matrices

<!-- grade_id: cell-c6274821bbcc2d34, points: 0, grade: false, solution: true -->
```python
# Create two matrices
# YOUR CODE HERE
raise NotImplementedError()

# Multiply both matrices using the @ operator
# YOUR CODE HERE
raise NotImplementedError()

# Multiply both matrices using the numpy
# YOUR CODE HERE
raise NotImplementedError()
```

## 3.3 Shapes

<!-- grade_id: cell-d24ae0e8eadb2299, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## 3.4 Load data from a file

<!-- grade_id: cell-f78f4a5d379fa752, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## 4. Matplotlib
Matplotlib is a comprehensive library for creating static, animated, and interactive visualizations in Python.

<!-- grade_id: cell-6e58671fe1441db6, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## 5. Pandas
Pandas is a fast, powerful, flexible and easy to use open source data analysis and manipulation tool, built on top of the Python programming language.

```python
import pandas as pd
```

<!-- grade_id: cell-cfe37d25fb6e3dfd, points: 0, grade: false, solution: true -->
```python
# Rad csv file
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-810fb9df9481543c, points: 0, grade: false, solution: true -->
```python
# List all column names
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f6071668179c0ec3, points: 0, grade: false, solution: true -->
```python
# Create new data frame
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-221e4d87c7a10caa, points: 0, grade: false, solution: true -->
```python
# Find outlier
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-453bc89e54d33ba6, points: 0, grade: false, solution: true -->
```python
# Remove outlier
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5862ecd7aad36e79, points: 0, grade: false, solution: true -->
```python
# Plot histogram
# YOUR CODE HERE
raise NotImplementedError()
```